# Exercise 1

An easy one to start with. Calculate the number of letters in your name:

<!-- grade_id: get_length_name, points: 0, grade: false, solution: true -->
```python
first_name = 'Xuefei'
last_name = 'Lu'

def get_length_name(first_name, last_name):
    length = 0
    
    ### BEGIN SOLUTION    
    length = len(first_name)+len(last_name)
    ### END SOLUTION
    
    return length

get_length_name(first_name,last_name)
```

Output:
```
8
```

Verify your answer (you typically get 1 answer for free, so you are aware of how the exercises are validated):

<!-- grade_id: get_length_name, points: 5, grade: true, solution: false -->
```python
assert get_length_name('Alexander','the Great') == 18

### BEGIN HIDDEN TESTS


assert get_length_name('johannes','de smedt') == 16


assert get_length_name('','') == 0


assert get_length_name('1','') == 1


assert get_length_name('','1') == 1
### END HIDDEN TESTS
### ''
#assert get_length_name('Alexander','the Great') == 17, 'something is wrong' # customise error message
```

# Exercise 2

The second exercise focuses on arithmetic and arrays.
Find the sum of the elements of two arrays.

<!-- grade_id: get_array_sums, points: 0, grade: false, solution: true -->
```python
def get_array_sums(arr1, arr2):
    sum1 = 0
    sum2 = 0
    
    ### BEGIN SOLUTION 
    
    # Loop first array and add each item to sum1
    for i in arr1:
        sum1 += i
    
    # Loop second array
    for i in arr2:
        sum2 += i
    
    ### END SOLUTION
    
    return sum1+sum2
```

```python
# or use sum([1, 2, 3, 4, 5])

def get_array_sums2(arr1, arr2):
    sum1 = sum(arr1)
    sum2 = sum(arr2)
    
    return sum1+sum2
```

Verify your answer:

<!-- grade_id: get_array_sums, points: 5, grade: true, solution: false -->
```python
assert get_array_sums([1, 2, 3, 4, 5],[6, 7, 8, 9, 10]) == 55

### BEGIN HIDDEN TESTS


assert get_array_sums([1, 2, 3, 4, 5],[6, 7, 8, 9, 10]) == 55


assert get_array_sums([],[6, 7, 8, 9, 10]) == 40


assert get_array_sums([1, 2, 3, 4, 5],[]) == 15


assert get_array_sums([],[]) == 0
### END HIDDEN TESTS
```

# Exercise 3

Calculate how many years you need to obtain a certain amount by starting with an initial amount and a given interest rate.

Don't forget: we use compound interest: you get interest on your interest.

<!-- grade_id: calculate_interest, points: 0, grade: false, solution: true -->
```python
def calculate_interest(amount, target, interest_rate):
    number_of_years = 0
    
    ### BEGIN SOLUTION    
    # Keep adding interest until you reach the target, or you go above it
    while amount < target:
        # Interest is added 
        amount = amount * (1 + interest_rate)
        
        # The number of years is raised
        number_of_years = number_of_years + 1
    ### END SOLUTION
    
    return number_of_years
```

Verify your answer:

<!-- grade_id: calculate_interest, points: 8, grade: true, solution: false -->
```python
assert calculate_interest(1000,2000,0.02) == 36

### BEGIN HIDDEN TESTS


assert calculate_interest(1000,1500,0.02) == 21


assert calculate_interest(1,1500,0.8) == 13


assert calculate_interest(1,80000000000,0.01) == 2524
### END HIDDEN TESTS
```

# Exercise 4

Remove the duplicates from a list.

<!-- grade_id: remove_duplicates, points: 0, grade: false, solution: true -->
```python
def remove_duplicates(input_list):
    output_list = []
    
    ### BEGIN SOLUTION    
    
    # Go through the list
    for item in input_list:
        
        # Check whether item is already in the output list, otherwise skip, as you would introduce duplicates again
        if item not in output_list:
            output_list.append(item)    ### If item is not in the output_list, add it into the list.
    ### END SOLUTION
    
    return output_list
```

Verify your answer:

<!-- grade_id: remove_duplicates, points: 4, grade: true, solution: false -->
```python
assert remove_duplicates(["Johannes", "Jamal", "Jamal", "Johannes", "Galina"]) == ['Johannes', 'Jamal', 'Galina']

### BEGIN HIDDEN TESTS


assert remove_duplicates([]) == []


assert remove_duplicates(["Johannes", "Galina"]) == ['Johannes', 'Galina']


assert remove_duplicates(["123", "345", "345"]) == ["123", "345"]
### END HIDDEN TESTS
```

# Exercise 5

Find the total number of points taught by a particular individual.

The input is two dictionaries: one containing the courses and their organisers, and one containing the courses and their points.

<!-- grade_id: cell-2aa20234cf44249b, points: 0, grade: false, solution: true -->
```python
### definition of dictionary: key and value
'''
thisdict = {
  "brand": "Ford",
  "model": "Mustang",
  "year": 1964
}
''' 
### two dictionaries: 1 - course with teacher: organiserDictionary, 2 - course with points: pointsDictionary

def find_points_for_organiser(givenOrganiser, organiserDictionary, pointsDictionary):
    total_points_givenOrganiser= 0
    
    ### BEGIN SOLUTION    
    # Loop the key-value pairs of the dictionary
    for course_in_loop, organiser_in_loop in organiserDictionary.items():     #.item: key and value: course:organiser_c    
        if organiser_in_loop == givenOrganiser:
            total_points_givenOrganiser += pointsDictionary[course_in_loop]   # points
    ### END SOLUTION
    
    return total_points_givenOrganiser
```

Verify your answer:

<!-- grade_id: find_points_for-organizer, points: 6, grade: true, solution: false -->
```python
courses = {"Analytics": "Johannes", "Mega analytics" : "Johannes", "Basic analytics" : "John"}
points = {"Analytics": 15, "Mega analytics": 10, "Basic analytics":5}

assert find_points_for_organiser('Johannes', courses, points) == 25

### BEGIN HIDDEN TESTS


assert find_points_for_organiser('John', courses, points) == 5


assert find_points_for_organiser('', courses, points) == 0
### END HIDDEN TESTS
```

```python

```