# Coding cross validation yourself

Now it is up to you to demonstrate your knowledge on the absenteeism dataset.

```python
import numpy as np
import pandas as pd


np.random.seed(40)

# First open absent.csv directly by double-click to have an overview of the data structure.
# And we see that we should let Jupyter read it every ";"
data = pd.read_csv("absent.csv",sep=';',index_col=0)

y = data['Absenteeism time in hours']
X = data.drop('Absenteeism time in hours',axis=1)

print(data.head)
```

Output:
```
<bound method NDFrame.head of     Reason for absence  Month of absence  Day of the week  Seasons  \
ID                                                                   
11                  26                 7                3        1   
36                   0                 7                3        1   
3                   23                 7                4        1   
7                    7                 7                5        1   
11                  23                 7                5        1   
..                 ...               ...              ...      ...   
11                  14                 7                3        1   
1                   11                 7                3        1   
4                    0                 0                3        1   
8                    0                 0                4        2   
35                   0                 0                6        3   

    Transportation expense  Distance from Residence to Work
```

## Applying cross validation

Apply cross-validation to this regression, where you
- **normalise** all variables,
- **calculate** the mean absolute error, mean squared error, and R-squared value
- and only return the test set value

(so here we have multiple metrics, careful about the function to choose)

<!-- grade_id: apply_cv, points: 0, grade: false, solution: true -->
```python

'''
Some hints:

You may want to define a function to call the process of cross validation.
The input will be your data and the number of fold.
The output object should be the results that contain metrics required as above. (score? cross_validate?)
    Here you can refer to the input parameter explanation of the function you choose to set True or False.
    You can consider input parameters: return_train_score and return_estimator to find how to define them.

A pipeline may be needed to contain normalizer along with classifier?

And don't forget to import library at the beginning.
'''

def apply_cv(X, y, nFold):    

# TODO: Here begins your answer.
    
    return
```

Validate your result:

<!-- grade_id: check_apply_cv, points: 5, grade: true, solution: false -->
```python
assert np.allclose(apply_cv(X,y,10)['test_neg_mean_absolute_error'],[-5.44169706, -4.88533949, -6.03776489, -6.53316679, -6.97119581,
       -5.6121378 , -6.20582349, -5.11806841, -5.58925382, -9.20383222])

assert np.allclose(apply_cv(X,y,10)['test_neg_mean_squared_error'],[ -57.35091985, -105.79371931,  -93.71342154, -120.49204514,
       -209.8988054 , -179.36848107, -180.39558093,  -90.82555524,
       -260.11583086, -309.20461226])

assert np.allclose(apply_cv(X,y,10)['test_r2'],[-0.11923377,  0.12709059,  0.09897902,  0.05605541,  0.1482223 ,
        0.07964105,  0.02008328,  0.14330421,  0.13328541,  0.06357028])

assert np.allclose(apply_cv(X,y,5)['test_neg_mean_absolute_error'],[-5.15813595, -6.30815503, -6.4512004 , -5.66281764, -7.53386985])

assert np.allclose(apply_cv(X,y,5)['test_neg_mean_squared_error'],[ -82.17300336, -112.00147052, -196.11580591, -135.58694837,
       -299.31974383])

assert np.allclose(apply_cv(X,y,5)['test_r2'],[0.05194651, 0.03335425, 0.11121968, 0.06534224, 0.052034  ])
```

Output:
```
TypeError                                 Traceback (most recent call last)
Input In [3], in <cell line: 1>()
----> 1 assert np.allclose(apply_cv(X,y,10)['test_neg_mean_absolute_error'],[-5.44169706, -4.88533949, -6.03776489, -6.53316679, -6.97119581,
      2        -5.6121378 , -6.20582349, -5.11806841, -5.58925382, -9.20383222])
      4 assert np.allclose(apply_cv(X,y,10)['test_neg_mean_squared_error'],[ -57.35091985, -105.79371931,  -93.71342154, -120.49204514,
      5        -209.8988054 , -179.36848107, -180.39558093,  -90.82555524,
      6        -260.11583086, -309.20461226])
      8 assert np.allclose(apply_cv(X,y,10)['test_r2'],[-0.11923377,  0.12709059,  0.09897902,  0.05605541,  0.1482223 ,
      9         0.07964105,  0.02008328,  0.14330421,  0.13328541,  0.06357028])

TypeError: 'NoneType' object is not subscriptable
```

## Cross validation with prediction

Calculate the root mean squared error of a prediction when it was in the test set:

<!-- grade_id: cell-2622132d0c9d2aed, points: 0, grade: false, solution: true -->
```python

'''
You can either define a function or directly calculate RMSE when data is given, but I would recommend to go with the function.

This function has your test set as input, as well as the number of fold.

Choose a function to predict, and don't forget that both normalizer and classifier are applied, so they should be contained together.

mean_squared_error from sklearn.metrics may be useful for the calculation of RMSE.
'''
    
def calculate_rmse(X, y, nFold):  

# TODO: Here begins your answer.

    return
```

Validate your result:

<!-- grade_id: cell-409719bb41aec1f7, points: 3, grade: true, solution: false -->
```python
assert np.allclose(calculate_rmse(X, y, 10),12.677377376954103)
assert np.allclose(calculate_rmse(X, y, 5),12.846765912010982)
assert np.allclose(calculate_rmse(X, y, len(X)),12.679949293998812)
```